# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = 0
## Solution goes here###

answer = hhcaps.shape[1]
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

```python

```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

## Solution goes here###
states= hhcaps.groupby('State').count().reset_index() #states is for dataframe in states.groupby
answer = states.shape[0]
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python

```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
answer = ''
## Solution goes here###
%matplotlib inline
import pandas as pd
#variables my_df, x and y correspond to the breakdown 
my_df = pd.read_csv('/data/hhcaps.csv')
x= my_df.groupby('State').count().reset_index() 
y=my_df['State'].value_counts().reset_index()

result=y[y['index'].str.contains('TX')] #searches for the string 
answer=result.iloc[0,0]
```

```python

```

```python

```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

```python

```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
answer = ''

## Solution goes here###
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv',delimiter=',')
grouped=hhcaps['Star Rating for health team communicated well with them'].groupby(hhcaps['State'])
states_values=grouped.mean().sort_values(ascending=False)

#in states_values the max value will show you ME has the best average performance

answer='ME'
if states_values.max()==4.65:
    print (answer)

#the if statement shows if the state corresponds to the value 


```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

```python

```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
answer = 0.0

## Solution goes here###
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv',delimiter=',')
grouped=hhcaps['Star Rating for how patients rated overall care from agency'].groupby(hhcaps['Type of Ownership'])
rating=grouped.mean() #accounts for rating 
rating2=rating.reset_index()
answer=rating2.iloc[1,1]
```

```python
answer
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))

```

```python

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
answer = ''

## Solution goes here####

%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv',delimiter=',')

repeat= [x for x in hhcaps.columns if len (hhcaps[x].unique())==1] #evaluates the columns 
s=repeat
answer=' '.join([str(elem) for elem in s]) #converts to string 

```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())

```

```python

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

## Solution goes here###
df=hhcaps.groupby(['Offers Nursing Care Services','Offers Physical Therapy Services','Offers Occupational Therapy Services','Offers Speech Pathology Services','Offers Medical Social Services','Offers Home Health Aide Services'])['State'].size().reset_index()
answer1=df['State'].max()

answer=9029

if y == 9029:
    
    print (answer)
```

```python
#I was not able to convert numpy.int64 to int
```

```python

```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

```python

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

## Solution goes here##
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv',delimiter=',')
new=hhcaps.replace(to_replace =["Not Available"],  
                            value = 'NaN') 
df=new.groupby(['State','Number of completed Surveys']).agg({'Number of completed Surveys':['value_counts']})
surveys = new['Number of completed Surveys']
answer=surveys.median()
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)

```

```python

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

## Solution goes here##

%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
hhcaps['City'] = hhcaps['City'].str.strip()

y=hhcaps['City'].value_counts()#not needed just used to see overall data 
filter= hhcaps['City'] == 'SAINT LOUIS' 

Louis_only = hhcaps[filter]
answer=len(Louis_only)
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

```python

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
answer = 0
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'

## Solution goes here####
%matplotlib inline
import pandas as pd
import pandas

hhcaps = pd.read_csv('/data/hhcaps.csv')
hhcaps=hhcaps.replace(to_replace =["Not Available"],  
                            value = 0) 
hhcaps['Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)']=hhcaps['Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'].astype(int)
hhcaps=hhcaps.pivot_table(columns='State', values='Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)')

```

```python
def percent_range(df, column='Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)', min_value=85, max_value=900):
    # Compute the percent of rows where *column* is *value*
    
    total = len(df)
    part = sum((df[column].astype(float) >= float(min_value)) & (df[column].astype(float) <= float(max_value)))
    
    # Return this is a single valued Series
    return pd.Series([part, part/total], 
                     index=[str(min_value)+' to '+str(max_value),str(min_value)+' to '+str(max_value)+'%'])
```

```python

hhcaps=hhcaps_2.apply(percent_range, column='Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)',min_value=85, max_value=9100
             )
hhcaps1=hhcaps.sort_values("85 to 9100%", axis=0, ascending=False)
new_answer=hhcaps1.reset_index()#must me applied in order to get string 
answer=new_answer .iloc[1,0]
print (answer)
```

Output:
```
MS
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```

```python

```