# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
# First, put in your user name below
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
MYTABLE='aavila' + '_data'

# Then, create a connection to the same database, using the 
# same user name and password in the `mysql-examples.ipynb` 
# file.  Use conn as the database connection variable.

import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd


host = 'slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2020'
password = 'SLUcor2020'
database = 'hds5210'




conn = None

```

```python

```

```python
conn = create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + '/' + 
                     database, echo=False)
```

```python

```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'paulboal_data')
assert(conn.connect())
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python

url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'

### SOLUTION GOES HERE####

df=pd.read_csv(url)
df.to_sql("Provider_Name", conn, if_exists='replace', index=False)
```

```python
dxyz=pd.read_sql_query("SELECT * FROM Provider_Name", conn) #I am sorry for creating additional files 
```

```python
dxyz
```

Output:
```
Provider_Name         NPI    CCN  \
0                            Sutter Bay Hospitals  1659439834  50008   
1    PRIME HEALTHCARE SERVICES - GARDEN GROVE LLC  1659538858  50230   
2                          ST MARY MEDICAL CENTER  1669456299  50300   
3                       MADERA COMMUNITY HOSPITAL  1669673646  50568   
4                    Temecula Valley Hospital Inc  1679816201  50775   
..                                            ...         ...    ...   
323                   KAISER FOUNDATION HOSPITALS  1326119967  50071   
324      COUNTY OF LOS ANGELES AUDITOR CONTROLLER  1336154020  50717   
325                    REEDLEY COMMUNITY HOSPITAL  1336167550  50192   
326  PROVIDENCE HEALTH SYSTEM-SOUTHERN CALIFORNIA  1336173269  50235   
327                   KAISER FOUNDATION HOSPITALS  1336294040  50411   

    Medicaid_EP_Hospital_Type           Street_Address           City  \
0        Acute Care Hospitals      CASTRO & DUBOCE STS  SAN FRANCISCO   
1        Acute Care Hospi
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (328,19)) 
assert(list(dxyz.columns) == [  'Provider_Name', 'NPI', 'CCN', 'Medicaid_EP_Hospital_Type', 'Street_Address' ,'City', 'County', 'State' ,'Zip_Code' ,'Payment_Year_Number','Program_Type', 'Total_payments', 'Last_Program_Year', 'Last_Payment_Year', 'Last_Payment_Criteria', 'Most_Recent_Disbursement_Amount','LONGITUDE', 'LATITUDE','FID'])
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Zip_Code` field and `population`'s `zipcode` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
answer = []

pd.read_sql_query(""" 
    SELECT Medicaid_EP_Hospital_Type, NPI, Zip_Code, total_payments, zipcode, population
    FROM
      NPI JOIN
      population ON Zip_Code=zipcode ORDER BY population
""",conn)

answer=['1194016923']

```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])
```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
answer = None

pd.read_sql_query(""" 
    SELECT Medicaid_EP_Hospital_Type, total_payments, zipcode, population, SUM(total_payments) as Total_Payments, SUM(population)
    FROM
      NPI JOIN
      population ON Zip_Code=zipcode
    group by Medicaid_EP_Hospital_Type, zipcode
    
""",conn)



```

Output:
```
Medicaid_EP_Hospital_Type  total_payments zipcode  population  \
0       Acute Care Hospitals         1292117   90004       62180   
1       Acute Care Hospitals         3840496   90012       31103   
2       Acute Care Hospitals         5580824   90015       18986   
3       Acute Care Hospitals         4195996   90017       23768   
4       Acute Care Hospitals         1553491   90018       49310   
..                       ...             ...     ...         ...   
73      Acute Care Hospitals         4279442   93534       39341   
74      Acute Care Hospitals         1017523   93551       50798   
75      Children's Hospitals          283784   90020       38967   
76      Children's Hospitals         6173289   90027       45151   
77      Children's Hospitals         8084174   90806       42399   

    Total_Payments  SUM(population)  
0        1292117.0          62180.0  
1        3840496.0          31103.0  
2        5580824.0          18986.0  
3        4195996.0          23768.
```

```python
pd.read_sql_query(""" 
    SELECT Medicaid_EP_Hospital_Type,total_payments, population, SUM(total_payments) as Total_Payments, SUM(population)
, (SUM(total_payments)/SUM(population))
 as Avg_Pay_per_Capita
    
    FROM
      NPI JOIN
      population ON Zip_Code=zipcode
   
       
        
    group by Medicaid_EP_Hospital_Type
    
""",conn)


```

Output:
```
Medicaid_EP_Hospital_Type  total_payments  population  Total_Payments  \
0      Acute Care Hospitals          494971       40311     243421965.0   
1      Children's Hospitals         8084174       42399      14541247.0   

   SUM(population)  Avg_Pay_per_Capita  
0        3336866.0             72.9493  
1         126517.0            114.9351
```

```python
pd.read_sql_query(""" 
    SELECT Medicaid_EP_Hospital_Type,total_payments, population, SUM(total_payments) as Total_Payments, SUM(population)
, (SUM(total_payments)/SUM(population)) as Avg_Pay_Per_Cap
    
    FROM
      NPI JOIN
      population ON Zip_Code=zipcode
   
       
        
    group by Medicaid_EP_Hospital_Type
    
""",conn)
```

Output:
```
Medicaid_EP_Hospital_Type  total_payments  population  Total_Payments  \
0      Acute Care Hospitals          494971       40311     243421965.0   
1      Children's Hospitals         8084174       42399      14541247.0   

   SUM(population)  Avg_Pay_Per_Cap  
0        3336866.0          72.9493  
1         126517.0         114.9351
```

```python
answer=pd.read_sql_query(""" 
    SELECT Medicaid_EP_Hospital_Type,total_payments, population, SUM(total_payments) as Total_Payments, SUM(population)
, (SUM(total_payments)/SUM(population)) as Avg_Pay_Per_Cap
    
    FROM
      NPI JOIN
      population ON Zip_Code=zipcode
   
       
        
    group by Medicaid_EP_Hospital_Type
    
""",conn)
```

```python

```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 17974.363)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 111.559)

```