# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 26.1 - When a Python program opens a file to read from it, no other programs can read from the same file at the same time.

> True or False


```python
answer261 = False 
```

### 26.2 - What is the advantage of opening files in Python using a `with` block?

> A) It ensures the file is closed as soon as you're done using it
>
> B) It runs faster than not using the `with` block
>
> C) It automatically loads the whole file into memory so that it can be accessed faster
>
> D) Actually, it is the only way to do it


```python
answer262 = "A"
```

### 26.3 - Which of the following is true?

> A) Many files only contain text, but many file formats do not use text
>
> B) All files contain only text, it's the just the text contents that explain the data to the computer
>
> C) All files contain only text in either JSON, XML, or CSV.
>
> D) Even if files contain only text, there still has to be a special code at the beginning to explain that this file will contain only text.


```python
answer263 = "B"
```

### 27.1 - Which of the following do you think is least likely to be used as a delimiter in structured files?

> A) `,`
>
> B) `|`
>
> C) `[tab]`
>
> D) `[space]`


```python
answer271 = "B"
```

### 27.2 - Which approaches can be used to read pipe-delimited (`|`) files?

*Report your answer as a list of options, e.g.: `["A","C"]`*

> A) Using the csv module
>
> B) Using the pandas module
>
> C) Using built-in file reading
>
> D) Using the psv module


```python
answer272 = ["A","B"]
```

### 27.3 - In a CSV file, how does the format change if you want to have a `,` in the actual contents of the file rather than as a delimiter?

> A) You can't do that.  You have to remove every occurence of `,`s
>
> B) Enclose the `,` using quotation marks
>
> C) Escape the `,` using a back-slash
>
> D) Nothing special, it will be obvious to the reader


```python
answer273 = "B"
```

### 28.1 - Select all of the data types that can be represented within a JSON format:

Provide your answer in the form of a list.  For example: `answer = ["A","B"]` 

> A) str
>
> B) list
>
> C) dict
>
> D) int


```python
answer281 = ["A"."C"]
```

### 28.2 - JSON stands for:

> A) Just a Simple Object Notation
>
> B) JavaScript Object Notation
>
> C) JavaScript Objects and Numbers
>
> D) Java Standard Object Notation


```python
answer282 = "B"
```

### 28.3 - Without just running the code, what will the following code print?

```
import json

data = """
[  {"id": { "last name": "Boal", "first name": "Paul", "middle initial": "E" }},
   {"id": { "last name": "Westhus", "first name": "Eric" }} ]
"""

names = json.loads(data)
list(names[1]['id'].keys())[1]
```

> A) 'Eric'
>
> B) 'first name'
>
> C) 'last name'
>
> D) 'Boal'


```python
answer283 = "B"
```

### 29.1 - In this class, XML Stands for:

> A) eXtensible Markup Language
>
> B) The 10th version of the Markup Language
>
> C) eXhange Markup Language
>
> D) X-Men Legends


```python
answer291 = "A"
```

### 29.2 - Which of the following is not a type of data in an XML document?

> A) tags
>
> B) attributes
>
> C) text
>
> D) list


```python
answer292 = "D"
```

### 29.3 - True or False, each element in an XML document can only have one child with a specific tag.

For example, is the following XML illegal because there are two child elements with the tag "kid":

```xml
<family>
  <parent role="dad">Paul Boal</parent>
  <kid role="daughter">Eleanor</kid>
  <kid role="daughter">Adelaide</kid>
</family>
```

Answer `False` if you think the above XML is valid.  

Answer `True` if you think the above XML is invalid.


```python
answer293 = True 
```

### 30.1 List of Allergies

In this directory, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
import json
def allergy_count(new_file):
    """

    This function takes a look into the file name allergies.json where the entry count is counted and evalauted based on the "type"
    of object is it. 
    The function allergy_count takes into account the information within the json file. A for loop is created and the output
    is a number and type of integer. 
    >>> allergy_count([new_total])
    3.0
    """
new_file=json.load(open('/data/allergies.json'))
for system in new_file:
    name=new_file.get("entry")
    new_total=len(name)
    

### END SOLUTION
```

```python
new_total
```

Output:
```
3
```

```python
type(new_total)
```

Output:
```
int
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count('/data/allergies.json')) == int
assert allergy_count('/data/allergies.json') == 3

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-242-aa09c0851708> in <module>
----> 1 assert type(allergy_count('/data/allergies.json')) == int
      2 assert allergy_count('/data/allergies.json') == 3

AssertionError:
```

```python

```

### 30.2 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
def allergy_list(json_file):
    """
    This function allergy_list will go through the entire json file to identify the patient's name, what he is allergic to and respective reaction.
    This function uses long for loops to get to desired output. There probably a shorter way to get around this but for me 
    it was easier to see the data loop around this huge file. The final output will give you a list of active_items with the name, 
    what he is allegric to and reaction. 
    >>> allergy_list('/data/allergies.json') 
    [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
    ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
    ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']] 
    """


### BEGIN SOLUTION

    
active_items=[] 

    #total_entry=0
reactions=json.load(open('/data/allergies.json'))
for reaction in reactions:
    new_one= reactions.get("entry")[0]
    for first in new_one:
        pat_0=new_one.get("resource")
        for name in pat_0:
            name1=pat_0.get("patient") 
            for display in name1:
                 final_name=name1.get("display")
            for same in pat_0:
                subs=pat_0.get("substance")
                for drug in subs:
                    drug1=subs.get("text")
                    rea=pat_0.get("reaction")[0]
                for new_react in rea:
                    final=rea.get("manifestation")[0]
                    for text in final:
                         final=final.get("text")
                            
for reaction in reactions:
    new_two= reactions.get("entry")[1]
    for first in new_two:
        pat_1=new_two.get("resource")
        for name_1 in pat_1:
            name2=pat_1.get("patient") 
            for display in name1:
                 final_name2=name2.get("display")
            for same2 in pat_1:
                subs2=pat_1.get("substance")
                for drug2 in subs2:
                    drug2=subs2.get("text")
                    rea2=pat_1.get("reaction")[0]
                for new_react2 in rea2:
                    final2=rea2.get("manifestation")[0]
                    for text2 in final2:
                         final_reaction=final2.get("text")
                            
                            

for reaction in reactions:
    new_three= reactions.get("entry")[2]
    for second in new_three:
        pat_2=new_three.get("resource")
        for name_2 in pat_2:
            name3=pat_2.get("patient") 
            for display1 in name3:
                 final_name3=name3.get("display")
            for same3 in pat_2:
                subs3=pat_2.get("substance")
                for drug3 in subs3:
                    drug3=subs3.get("text")
                    rea3=pat_2.get("reaction")[0]
                for new_react3 in rea3:
                    final3=rea3.get("manifestation")[0]
                    for text3 in final3:
                         final_reaction3=final3.get("text")                            
                            

                            
active_items.append([final_name,drug1,final])
active_items.append([final_name2,drug2,final_reaction])
active_items.append([final_name3,drug3,final_reaction3])
```

```python
active_items
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list('/data/allergies.json') == output

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-6f71eb1cb492> in <module>
      3  ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
      4 
----> 5 assert allergy_list('/data/allergies.json') == output

AssertionError:
```

### 30.3 Allergy Reaction

Write a function called `allergy_reaction(json_file,substance)` that takes two parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json
      
### BEGIN SOLUTION
def allergy_list(json_file):
    def allergy_reaction(name,results):
        """
        This function calls another function allergy_list and loops through the list active_list. Based on the previous 
        problem, you can see whether or not certain strings are within the list. Each for loop evalautes the corresponding strings.
        If none of the strings match the result will be zero. 
        >>>allergy_reaction('allergies.json','STRAWBERRY')
        Anaphylaxis
        """
    for x in active_items[0]:
        if x=='PENICILLIN G':
            print('Hives')
    for x in active_items[1]:
        if x=='SHELLFISH-DERIVED PRODUCTS':
            print('Itching')
    for x in active_items[2]:
        if x=='STRAWBERRY':
            print('Anapylaxis')
    for x in active_items[0]:
        if x!='PENICILLIN G' or x!='SHELLFISH-DERIVED PRODUCTS'or x!='STRAWBERRY':
            print(None)
            break
            
### END SOLUTION
```

```python

```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction('allergies.json','PENICILLIN G') == 'Hives'
assert allergy_reaction('allergies.json','SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction('allergies.json','STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction('allergies.json','PENICILLIN') == None
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-22-4b2f3d62cdfa> in <module>
----> 1 assert allergy_reaction('allergies.json','PENICILLIN G') == 'Hives'
      2 assert allergy_reaction('allergies.json','SHELLFISH-DERIVED PRODUCTS') == 'Itching'
      3 assert allergy_reaction('allergies.json','STRAWBERRY') == 'Anaphylaxis'
      4 assert allergy_reaction('allergies.json','PENICILLIN') == None

NameError: name 'allergy_reaction' is not defined
```

### STRETCH - For those looking for an additional challenge

The Coronavirus is creating quite the stir right now.  There are some sources suggesting that trends show it is going to be significantly more serious than SARS was back in the 2002 timeframe.  Here's one visualization trying to demonstrate that: https://www.reddit.com/r/China_Flu/comments/ev2b4v/i_updated_some_charts_comparing_this_outbreak/

Someone on Kaggle has generously already compiled a dataset based on information from Johns Hopkins about the Coronavirus outbreak.  https://www.kaggle.com/brendaso/2019-coronavirus-dataset-01212020-01262020  Create a Kaggle account, if you don't already have one.  Download this data set and then upload it to your Jupyter Home folder.  (The "up arrow" button is for uploading a file.)

Use Python's built-in `csv` module to read the data from this file and generate the following information: **what are the total confirmed cases in all of Mainland China as of the latest information in the data set?**  Some important things to note:
* Each entry for a given city has the **cumulative** number of cases.  So that column is not additive (it cannot be summed).  You'll have to find a way to filter your data for the last day for each city, then total those up.
* If you choose to parse the date column, you will want to lookup how to do that using Python's `datetime` module.  Especially the `strptime` function.  https://docs.python.org/3/library/datetime.html#strftime-strptime-behavior  Hint: you can parse a date string in the format 2/17/2020 using the code below.  This link will tell you what things like `%m` and `%Y` mean.

```
from datetime import datetime
d = datetime.strptime('2/17/2020', '%m/%d/%Y')
```

If you want to take this another step, **create a list of tuples that contain (observate date, total confirmed) totalled over all locations represented in the data**

```python

```

```python

```